# Assignment 2

More on Accumulation

Write code that **uses `map`** to assign to the variable `map_testing` all the elements from `lst_check`, prepending the string `"Fruit: "` to the beginning of each element. *Note: be sure to use `list` to convert the result from `map` into a list.*

<!-- grade_id: cell-33f705472409f897, points: 0, grade: false, solution: true -->
```python
lst_check = ['plums', 'watermelon', 'kiwi', 'strawberries', 'blueberries', 'peaches', 'apples', 'mangos', 'papaya']

map_testing = list(map(lambda x: f"Fruit: {x}", lst_check))

print(map_testing)

```

Output:
```
['Fruit: plums', 'Fruit: watermelon', 'Fruit: kiwi', 'Fruit: strawberries', 'Fruit: blueberries', 'Fruit: peaches', 'Fruit: apples', 'Fruit: mangos', 'Fruit: papaya']
```

<!-- grade_id: cell-9270b9c7ad936d05, points: 1, grade: true, solution: false -->
```python
assert 'Fruit: plums' in map_testing
assert 'Fruit: kiwi' in map_testing

```

Below, we have provided a list of strings called `countries`. **Use `filter`** to produce a list called `b_countries` that only contains the strings from countries that begin with `'B'`. *Note: be sure to use `list` to convert the result from `filter` into a list.*

<!-- grade_id: cell-728dbbed11e70922, points: 0, grade: false, solution: true -->
```python
countries = ['Canada', 'Mexico', 'Brazil', 'Chile', 'Denmark', 'Botswana', 'Spain', 'Britain', 'Portugal', 'Russia', 'Thailand', 'Bangladesh', 'Nigeria', 'Argentina', 'Belarus', 'Laos', 'Australia', 'Panama', 'Egypt', 'Morocco', 'Switzerland', 'Belgium']

b_countries = list(filter(lambda strn: strn[0] == 'B', countries))

print(b_countries)
```

Output:
```
['Brazil', 'Botswana', 'Britain', 'Bangladesh', 'Belarus', 'Belgium']
```

<!-- grade_id: cell-a3a6ce379963c355, points: 1, grade: true, solution: false -->
```python
assert 'Brazil' in b_countries
assert 'Canada' not in b_countries

```

Below, we have provided a list of tuples that contain the names of Game of Thrones characters. The first item in each tuple is their last (family) name. The second item is their first (given) name. Using a **list comprehension**, create a list of strings called `first_names` that contains only the first names of everyone in the original list.

<!-- grade_id: cell-6d4a02615ad9dc17, points: 0, grade: false, solution: true -->
```python
people = [('Snow', 'Jon'), ('Lannister', 'Cersei'), ('Stark', 'Arya'), ('Stark', 'Robb'), ('Lannister', 'Jamie'), ('Targaryen', 'Daenerys'), ('Stark', 'Sansa'), ('Tyrell', 'Margaery'), ('Stark', 'Eddard'), ('Lannister', 'Tyrion'), ('Baratheon', 'Joffrey'), ('Bolton', 'Ramsey'), ('Baelish', 'Peter')]
first_names=[name[1] for name in people]
print(first_names)
```

Output:
```
['Jon', 'Cersei', 'Arya', 'Robb', 'Jamie', 'Daenerys', 'Sansa', 'Margaery', 'Eddard', 'Tyrion', 'Joffrey', 'Ramsey', 'Peter']
```

<!-- grade_id: cell-1d94c540a0067985, points: 1, grade: true, solution: false -->
```python
assert 'Jon'  in first_names
assert 'Snow' not in first_names

```

Use a **list comprehension** to create a list called `lst2` that "doubles" each element in the list, `lst`. Use `* 2` to "double" numbers (`5 * 2 == 10`), strings (`"xyz" * 2 == "xyzxyz"`), and lists (`['a', 'b'] * 2 == ['a', 'b', 'a', 'b']`).

<!-- grade_id: cell-3d3befbc2c890e0a, points: 0, grade: false, solution: true -->
```python
lst = [["hi", "bye"], "hello", "goodbye", [9, 2], 4]
lst2 =[element*2 for element in lst]
print(lst2)
```

Output:
```
[['hi', 'bye', 'hi', 'bye'], 'hellohello', 'goodbyegoodbye', [9, 2, 9, 2], 8]
```

<!-- grade_id: cell-f1ef42d126ceaad3, points: 1, grade: true, solution: false -->
```python
assert "hellohello" in lst2

```

Below, we have provided a list of tuples that contain students' names and their final grades in PYTHON 101. Using a **list comprehension**, create a new list `passed` that contains the *names* of students who passed the class (had a final grade of `70` or greater). `passed` should be a list of strings (names without final grades) in the same order as `students`.

<!-- grade_id: cell-0026cd972495389b, points: 0, grade: false, solution: true -->
```python
students = [('Tamara', 95), ('Linda', 63), ('Carl', 70), ('Bob', 100), ('Raymond', 50), ('Sue', 75)]
passed=[num[0] for num in students if num[1]>=70]
print(passed)
```

Output:
```
['Tamara', 'Carl', 'Bob', 'Sue']
```

<!-- grade_id: cell-7365646af3677b07, points: 1, grade: true, solution: false -->
```python
assert 'Tamara' in passed
assert 'Linda' not in passed

```

Write code using `zip` and `filter` so that these lists (`l1` and `l2`) are combined into one big list of pairs and assigned to the variable `opposites` if both elements of the pair are longer than `3` characters. *Note: be sure to use `list` to convert the result into a list.*

<!-- grade_id: cell-4a5785241c831a5a, points: 0, grade: false, solution: true -->
```python
l1 = ['left', 'up', 'front']
l2 = ['right', 'down', 'back']
l3 = zip(l1, l2)
print(l3)
opposites = list(filter(lambda s: len(s[0])>3 and len(s[1])>3, l3))
print(opposites)
```

Output:
```
<zip object at 0x7956e6ce1a40>
[('left', 'right'), ('front', 'back')]
```

<!-- grade_id: cell-077fa8cd007e9e9c, points: 1, grade: true, solution: false -->
```python
assert ('left', 'right') in opposites
# omit if either word is shorter than 3 characters...
assert ('up', 'down') not in opposites

```

Below, we have provided a `species` list and a `population` list. **Use `zip`** to combine these lists into one list of tuples called `pop_info`. From this list, create a new list called `endangered` that contains the names of species whose populations are below `2500`.

<!-- grade_id: cell-5fcdbdf62ecae78e, points: 0, grade: false, solution: true -->
```python
species = ['golden retriever', 'white tailed deer', 'black rhino', 'brown squirrel', 'field mouse', 'orangutan', 'sumatran elephant', 'rainbow trout', 'black bear', 'blue whale', 'water moccasin', 'giant panda', 'green turtle', 'blue jay', 'japanese beetle']

population = [10000, 90000, 1000, 2000000, 500000, 500, 1200, 8000, 12000, 2300, 7500, 100, 1800, 9500, 125000]
pop_info = list(zip(species,population))
print(pop_info)
endangered = [name[0] for name in pop_info if name[1]<2500]
print(endangered)
```

Output:
```
[('golden retriever', 10000), ('white tailed deer', 90000), ('black rhino', 1000), ('brown squirrel', 2000000), ('field mouse', 500000), ('orangutan', 500), ('sumatran elephant', 1200), ('rainbow trout', 8000), ('black bear', 12000), ('blue whale', 2300), ('water moccasin', 7500), ('giant panda', 100), ('green turtle', 1800), ('blue jay', 9500), ('japanese beetle', 125000)]
['black rhino', 'orangutan', 'sumatran elephant', 'blue whale', 'giant panda', 'green turtle']
```

<!-- grade_id: cell-78396308d9912e18, points: 1, grade: true, solution: false -->
```python
assert ('golden retriever', 10000) in pop_info
assert 'golden retriever' not in endangered
assert 'black rhino' in endangered

```